## Exercise 06.1 (selecting and passing data structures)

The task in Exercise 04 for computing the area of a triangle involved a function with six arguments ($x$ and $y$ components of each vertex). With six arguments, the likelihood of a user passing arguments in the wrong order is high.

Use an appropriate data structure, e.g. a `list`, `tuple`, `dict`, etc, to develop a new version of the function with an improved interface (the interface is the arguments that are passed to the function) with fewer arguments. Add appropriate checks inside your function on the input data.

<!-- grade_id: cell-50c4566d7bcbffaf, points: 0, grade: false, solution: true -->
```python
c1 = [0, 0]
c2 = [0, 2]
c3 = [3, 0]

c1_x = c1[0]
c1_y = c1[1]
c2_x = c2[0]
c2_y = c2[1]
c3_x = c3[0]
c3_y = c3[1]

def area(c1_x, c1_y, c2_x, c2_y, c3_x, c3_y):
  return abs((c1_x*(c2_y-c3_y) + c2_x*(c3_y-c1_y) + c3_x*(c1_y-c2_y))*1/2);
print(area(c1_x, c1_y, c2_x, c2_y, c3_x, c3_y))
```

Output:
```
3.0
```

## Exercise 06.2 (selecting data structures)

For a simple (non-intersecting) polygon with $n$ vertices, $(x_0, y_0)$, $(x_1, y_1)$, . . , $(x_{n-1}, y_{n-1})$, the area $A$ is given by
$$
A = \left| \frac{1}{2} \sum_{i=0}^{n-1} \left(x_{i} y _{i+1} - x_{i+1} y_{i} \right) \right|
$$
and where $(x_n, y_n) = (x_0, y_0)$. The vertices should be ordered as you move around the polygon.

Write a function that, given only the vertex coordinates, computes the area of a simple polygon. Test your function for some simple shapes. Pay careful attention to the range of any loops. **Test for polygons without a vertex at the origin.**

<!-- grade_id: cell-4b5a619405787746, points: 0, grade: false, solution: true -->
```python
import math

#c1 = [0, 0, 0, 2, 3, 0]
#c2 = [1, 3, 5, 7, 9, 14]
c3 = [1, 2, 45, 65, 13, 27, 11, 15, 17, 92, 13, 12, 12, 47]
def extract_x(coordinates):
  x_values = []
  for i in range(0, len(coordinates), 2):
    x_values.append(coordinates[i])
  return x_values;
def extract_y(coordinates):
  y_values = []
  for i in range(1, len(coordinates), 2):
    y_values.append(coordinates[i])
  return y_values;
def area(x_values, y_values):
  if len(x_values) + len(y_values) >=6:
    n = len(x_values)
    area = 0
    for i in range(n-1):
      area += (x_values[i] * y_values[(i + 1)]) - (x_values[(i + 1)] * y_values[i])
    area += (x_values[-1] * y_values[0]) - (x_values[0] * y_values[-1])
    return abs(area/2);
  else:
    print("Invalid Input")
print(area(extract_x(c3), extract_y(c3)))


```

Output:
```
226.0
```

## Exercise 06.3 (indexing)

Write a function that uses list indexing to add two vectors of arbitrary length, and returns the new vector. Include a check that the vector sizes match. Print a warning message if there is a size mismatch. The more error information you provide, the easier it would be for someone using your function to debug their code.

Add some tests of your code.

#### Hint: You can create a list of zeros of length `n` by

    z = [0]*n

<!-- grade_id: cell-47daca03f7a96b62, points: 0, grade: false, solution: true -->
```python
def sum_vector(x, y):
    """Return sum of two vectors"""
    if len(x) != len(y):
        print("Warning: Vectors are not the same length. Vector 1 has length" , len(x) , "and Vector 2 has length" , len(y) , "so they can't be added together")
    elif len(x) == len(y):
        result = [x[i] + y[i] for i in range(len(x))]
        return result;
x = [0, 4.3, -5, 7]
y = [-2, 7, -15, 1]
print(sum_vector(x, y))

```

Output:
```
[-2, 11.3, -20, 8]
```

<!-- grade_id: cell-95f5819810d2b2c7, points: 0, grade: true, solution: false -->
```python
# tests ##
a = [0, 4.3, -5, 7]
b = [-2, 7, -15, 1]
c = sum_vector(a, b)
assert c == [-2, 11.3, -20, 8]
```

### Optional (advanced)

Try writing a one-line version that performs the `sum_vector` operation using list comprehension and the built-in function [`zip`](https://docs.python.org/3/library/functions.html#zip).

<!-- grade_id: cell-c182b3f0d8952692, points: 0, grade: true, solution: true -->
```python
...
```

## Exercise 06.4 (dictionaries)

Create a dictionary that maps college names (the key) to college abbreviations for at least 5 colleges
(you can find abbreviations at https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges).
From the dictionary, produce and print

1. A Python dictionary from college abbreviation to name; and
1. A Python list of college abbreviations sorted into alphabetical order.

*Optional extension:* Create a dictionary that maps college names (the key) to dictionaries of:

- College abbreviation
- Year of foundation
- Total number students

for at least 5 colleges. Take the data from https://en.wikipedia.org/wiki/Colleges_of_the_University_of_Cambridge#Colleges. Using this dictionary,

1. Find the college with the greatest number of students and print the abbreviation; and
2. Find the oldest college, and print the number of students and the abbreviation for this college.

<!-- grade_id: cell-bc271dc3885e19c5, points: 0, grade: false, solution: true -->
```python
colleges = {
    "CHU": "Churchill College",
    "CL": "Clare College",
    "JE": "Jesus College",
    "CHR": "Christ's College",
    "DOW": "Downing College",
    "EM": "Emmanuel College",
    "M": "Magdalene College",
    "CAI": "Gonville and Caius College"
}

sorted_colleges = dict(sorted(colleges.items()))
print(sorted_colleges)
```

Output:
```
{'CAI': 'Gonville and Caius College', 'CHR': "Christ's College", 'CHU': 'Churchill College', 'CL': 'Clare College', 'DOW': 'Downing College', 'EM': 'Emmanuel College', 'JE': 'Jesus College', 'M': 'Magdalene College'}
```

#### Optional extension

<!-- grade_id: cell-7040c4053c72ce09, points: 0, grade: false, solution: true -->
```python
...
```